## Topographische und hydrographische Modelle (120.109)

### Protokoll 3

#### Aufgabe 3 (4 Punkte)

Erstellen Sie in QGIS jeweils eine Wassertiefenkarte für die Datensätze aus Aufgabe 1 (Laserbathymetrie) und Aufgabe 2 (Photobathymetrie)

- Verwenden Sie als Grundlage das bereitgestellte Orthophoto (~/shared/Daten/Protokoll_3/orthophoto.tif)
- Verwenden Sie für die Visualisierung und Erstellung Ihrer Karten die mit OPALS erstellten Tiefenlinien (*...wtiefe_iso.shp*) und 
- und die farbkodierten Tiefenstufenkarten (*..._wtiefe_zco.tif*) 
- Fügen sie im Drucklayout Titel, die entsprechende Farblegende und einen Maßstabsrahmen (Grid wie in Protokoll 1 und 2) hinzu.
- Speichern Sie Ihre erstellten Karten als *wassertiefenkarte_laser.pdf* bzw. *wassertiefenkarte_photo.pdf*  ab. 

Vergessen Sie nicht darauf, das Koordinatensystem vor dem Datenimport in QGIS einzustellen (ETRS89/UTM33, EPSG Code: 25833).

<!-- grade_id: p3_3a_test1, points: 1, grade: true, solution: false -->
```python
# Sichtbare Überprüfungen
from os.path import exists
assert exists('wassertiefenkarte_laser.pdf')
assert exists('wassertiefenkarte_photo.pdf')
```

### Manuelle Überprüfung des Karteninhalts
- Überprüfung der Wassertiefenkarten (Farbkodierung, Isolinien)
- Überprüfung von Titel, Legende und Kartenrahmen

### Analyse

Analysieren Sie die Unterschiede der beiden Tiefenkarten. Welches der beiden Modelle liefert konsistentere Ergebnisse? Wo stimmen die beiden Modelle gut überein, wo weniger gut? Was könnten Gründe für Abweichungen sein?

<!-- grade_id: p3_aufgabe_3b, points: 2, grade: true, solution: true -->
Die konsistenteren Ergebnisse liefert eindeutig die Photobathymetrie. Die Auswertung zeigt sehr glatte Übergänge und realistischere Wassertiefen für einen Fluss dieser Größe als es die der Laserbathymetrie tut. Bei letzterer scheint die Pielach an der tiefsten Stelle nicht einmal einen Meter tief zu sein. Außerdem erscheint der Gewässerboden sehr löchrig ohne geglättete Übergänge, was für ein Fließgewässer ungewöhnlich wäre.

Trotzdem gibt es auch Gemeinsamkeiten. Die innere Schotterbank wird z.B. auf beiden Karten sehr ähnlich dargestellt und auch der grundsätzliche Verlauf des Flusses inkl. der Positionen der tiefsten Stellen wird sowohl photo- als auch laserbathymetrisch gleichermaßen ermittelt.

Die Gründe für die Abweichungen der Laserbathymetrie bei der Wassertiefe könnten einerseits ein schlechtes Durchdringen des Wassers durch den Laserstrahl beinhalten (wobei normalerweise der Laser sogar tiefer eindringt als das Sonnenlicht reflektiert wird), andererseits könnte die schlechtere Lagegenauigkeit der Laservermessung dazu geführt haben, dass schmale tiefe Rinnen, wie sie in der gegebenen Schlinge der Pielach vorkommen, nicht ausreichend genau berechnet werden können bzw. durch Unsicherheitsstatistik ausgeflacht werden. Schließlich könnte auch der Umstand eine Rolle spielen, dass Photobathymetrie deutlich kostengünstiger und einfacher durchzuführen ist und deshalb häufigere und umfassendere Befliegungen stattfinden können, als dies beim Laserscanning finanzierbar wäre. Das führt dazu, dass deutlich dichtere Punktewolken aus der Photobathymetrie entstehen, die wiederum eine genauere Rekonstruktion der Topografie erlauben.

#### Bonusaufgabe
Erstellen Sie eine Karte mit den Differenzen der aus Laser- bzw. Photobathymetrie ermittelten Wassertiefen und diskutieren Sie die Abweichungen in kurzen Worten. Je nach Qualität der Abbildung und Antwort werden bis zu 2 Bonuspunkte vergeben.

<!-- grade_id: cell-bcde4579051cb104, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## !!! Aufräumen vor dem Abggeben !!!

Bevor Sie Protokoll_3 (via Submit) abgeben, stellen Sie sicher, dass alle Zwischen- und Resultatdateien ausgenommen der in QGIS erstellten PDF-Karten entfernt wurden. Führen Sie dazu die Zellen *Sicherstellung der Nicht-Existenz der Ergbnisdateien* in Aufgabe 1 und 2 aus. Im (Protokoll_3) Verzeichnis sollten sich dann folgende Dateien befinden:

    Aufgabe_1.ipynb
    Aufgabe_2.ipynb
    Aufgabe_3.ipynb
    opals.cfg
    pcl-laserbathy.png
    pcl-photobathy.png
    wassertiefenkarte_laser.pdf
    wassertiefenkarte_photo.pdf
